```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

# CE 93: Lab Assignment 05

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to learn about properties of continuous random variables and some of the common discrete distributions. We will not load any data. Instead, we will create the random variables ourselves.

## Instructions 
**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and these should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

import math
import numpy as np
import pandas as pd
import statistics as stats
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import resources
from scipy import integrate
from scipy.stats import * 
```

### Properties of Probability Density Function 

Let's first review some basic concepts.

<font color='red'>**Question 1.0.**</font> What is(are) the property(ies) of a probability density function? 
 Assign ALL that apply to the variable `q1_0`. (0.5 pts)

**A**. The probability density function is a function that assigns probabilities to each possible outcome of a continuous random variable \
**B**. The probability density function can take values greater than 1 as long as the integral over the entire range is equal to 1 \
**C**. The probability density function can take negative values as long as the integral over the entire range is equal to 1 \
**D**. To obtain probabilities from a probability density function, we need to integrate it \
**E**. The probability of a continuous random variable being exactly equal to any outcome is always 0 \
**F**. The probability of a continuous random variable being exactly equal to any outcome is always positive

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q1_0 = ...
q1_0
```

```python
grader.check("q1.0")
```

### Defining Functions in Python

Before we define our own probability density functions, let's review how to define functions in Python.

In programming, a function is a named sequence of statements that takes input data and returns an output. Similar to mathematical functions, Python functions take in data, perform on them a sequence of statements, and then return an output.

An example of a mathematical function is: $f(x)=1.25(1-x^4)$

In Python, a function that serves the same task looks like:

```python
def PDF_x(x):
    return 1.25 * (1 - x**4)
```

We will focus on defining simple functions. The most common way to define a function in Python is using the keyword `def` (short for define), as shown below:

```python
def function_name(argument_1, argument_2, ...):
    
    return output_parameters

```

The syntax includes:
1. Function header that begins with the keyword `def`, which is used to define a function
2. `function_name`, which is the identifier chosen by the programmer (you)
> Functions must conform to a naming scheme similar to variables. They can only contain alphanumeric characters and underscores, and the first character must be a letter or an underscore. Also, function names are case-sensitive.
3. Input arguments `(argument_1, argument_2, ...)`, which are a comma-separated list of inputs that will be sent to the function 
4. `:` colon, which is used at the end of the function header, after the input arguments
5. `return`, which will return (send back) the value(s) after the keyword to the calling statement

### Calling Functions in Python

Defining a function does not make it run. To run or use a function, we need to **call** it. To call a function, use its name followed by the list of input arguments between parentheses: `function_name(argument_1, argument_2, ...)`.

For example, if we defined the following function:

```python
def PDF_x(x):
    return 1.25 * (1 - x**4)
```

Then, if we want to evaluate the probability density function for $x = 0.5$, we can use:

```python
>>> PDF_x(0.5)

1.171875
```

### Defining Probability Density Function

Let $Z$ be a continuous random variable with the following probability density function:

$$ f_Z(z)=   \left\{
\begin{array}{ll}
      \frac{2}{30} (1+\frac{z}{10}) & 0\leq z \leq 10\\
      0 & \text{otherwise}\\
\end{array} 
\right.  $$

<font color='red'>**Question 2.0.**</font> Complete the Python function below, `PDF_z(z)`, which takes as input a scalar number or a numeric sequence (i.e., list, tuple, NumPy array) of numeric values and returns the density(ies) at the input value(s) using the above equation for $f_Z(z)$. (0.5 pts)

Assume that the input values are $0\leq z \leq 10$, so only return the input(s) evaluated using the equation $\dfrac{2}{30} \left(1+\dfrac{z}{10}\right)$.

Then try out your new `PDF_z()` function using `z = 3` and assign the output to `q2_0`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable.

```python
# ANSWER CELL

# define a function to evaluate the PDF
def PDF_z(z):
    return ...

# Get the density for z = 3
q2_0 = ...
print(f'f(z = 3) = {q2_0}')
```

```python
grader.check("q2.0")
```

We can also define an array with many $z$ values and evaluate the PDF at these values using the function `PDF_z()`. Because $Z$ is a continuous random variable with $0\leq z \leq 10$, we can create an array with many values between 0 and 10. We can create an array with **a certain number of evenly spaced points between a start and end value** using `np.linspace()` function in NumPy:

```python
np.linspace(start, stop, num)
```

where:

* `start`: a number specifying at which value to start the array. Here, this is a required argument.
* `stop`: a number specifying at which value to stop the array. This is by default **included** in the array, unlike `np.arange()`. This is a required argument.
* `num`: an integer number specifying how many values to generate between `start` and `stop`, including `start` and `stop`. This is an optional argument, and thus, if not specified, it will take the default value 50.

**Example:**

To create $10$ equally spaces numbers within the interval $[0, 1]$, we can use:
```python
>>> x = np.linspace(0, 1, 10)
```

<font color='red'>**Question 2.1.**</font> Create an array named `z` with 1000 evenly spaced values between 0 and 10 (range of z). Then, use the function that you defined to calculate the PDF for every value in the array `z`. Assign the output to `PDF_z_values`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.5 pts)

```python
# ANSWER CELL

# create an array z with 1000 evenly spaced values between 0 and 10 (range of z)
z = ...

# use the function you defined to calculate the PDF for every value in the array z
PDF_z_values = ...

print(f'PDF values: [{PDF_z_values[0]}, {PDF_z_values[1]}, ..., {PDF_z_values[-2]}, {PDF_z_values[-1]}]' if not isinstance(PDF_z_values, type(Ellipsis)) else None)
```

```python
grader.check("q2.1")
```

### Calculating Probabilities

To calculate probabilities from the PDF of a continuous random variable, we have to integrate:

$$P(a \leq x \leq b) = \int_{a}^{b} f(x) \, dx$$

We can use the `integrate.quad()` function to perform integration directly in Python as follows:

```python
integrate.quad(func, a, b)
```

where:

* `func`: the name of the Python function you want to integrate
* `a`: a number specifying the lower limit of integration
* `b`: a number specifying the upper limit of integration

Note that `integrate.quad()` returns a tuple with two values, in the following order:
1. The integral of `func` from `a` to `b`
2. An estimate of the absolute error in the result

We will only focus on the value of the integral. So we can use indexing to get the integral using `integrate.quad(func, a, b)[0]`, where `[0]` is simply indexing the first element, which is the value of the integral.

**Example:**

To calculate $P(1 \leq z \leq 5)$, we can use:
```python
>>> integrate.quad(PDF_z, 1, 5)[0]

0.34666666666666673
```

<font color='red'>**Question 2.2.**</font> Calculate the following probabilities: (0.5 pts)

* Calculate $P(0 \leq z \leq 3)$ and assign the output to `P_0_3`
* Calculate $P(5 \leq z \leq 9)$ and assign the output to `P_5_9`


Do not just manually type the numeric answers. Use Python expressions that return the desired answer and assign the expressions to the variable.

```python
# ANSWER CELL

# calculate P(0 < z < 3)
P_0_3 = ...
print(f'P(0 < z < 3) = {P_0_3:.3f}' if not isinstance(P_0_3, type(Ellipsis)) else None)

# calculate P(5 < z < 9)
P_5_9 = ...
print(f'P(5 < z < 9) = {P_5_9:.3f}' if not isinstance(P_5_9, type(Ellipsis)) else None)
```

```python
grader.check("q2.2")
```

So now, we know how to calculate the PDF of any continuous function and evaluate its integral easily in Python.

In the lecture, we defined two criteria that should be satisfied for a function to be a valid PDF:

$$1. \quad f(x) \geq 0 \text{ for all }x$$
$$2. \quad \displaystyle \int \limits_{-\infty} ^{+\infty} f(x)dx=1$$

We can check these conditions for the random variable $Z$ that we have been working with using:
1. `np.all(PDF_z_values >= 0)`
> If this returns `True`, then $f(z) \geq 0$ for all $z$
2. `round(integrate.quad(PDF_z, 0, 10)[0], 3) == 1`
> If this returns `True`, then the PDF integrates to 1. **Note that we are integrating the PDF here from 0 to 10, because this is the range over which $z$ is defined.** We are also rounding the answer due to numerical errors.

Both conditions should be satisfied for a PDF to be valid. We can check that both conditions are satisfied using logical operators: 

`(np.all(PDF_z_values >= 0)) and (round(integrate.quad(PDF_z, 0, 10)[0], 3) == 1)`

If both conditions are satisfied, this will return `True`. Otherwise, it will return `False`.

Run the code cell below to check if the function in `PDF_z` is a valid PDF.

```python
# check if f(z) is a valid PDF

(np.all(PDF_z_values >= 0)) and (round(integrate.quad(PDF_z, 0, 10)[0], 3) == 1)
```

### Valid PDFs

Next, you will use the method from above to check if a PDF is valid.

<font color='red'>**Question 3.0.**</font> Using the methods we used above to check if a PDF is valid, is the following function a valid PDF? You need to perform all calculations in `Python` rather than by hand. (1 pt)

$$ f_W(w)=   \left\{
\begin{array}{ll}
      \frac{3}{4}(1-w^2) & 0 \leq w \leq 1\\
      0 & \text{otherwise}\\
\end{array} 
\right.  $$

Follow these steps:
1. Define a function that calculates the PDF values of $W$. Name the function `PDF_w`.
2. Define an array `w` that includes 1000 evenly spaced values over the range for which $W$ is defined.
3. Calculate the PDF for every value in the array `w`. Assign the output to `PDF_w_values`.
4. Integrate $f_W(w)$ over the range for which $W$ is defined. Assign the integral value to `integral`.
5. Finally, check if both conditions are satisfied for it be a valid PDF.


Once you've completed all of the above, answer the following question: Is $f_W(w)$ a valid PDF? Assign your answer to the variable `q3` as a string.

**A.** Yes it is a valid PDF.  \
**B.** No it is not a valid PDF because the integral over all possible values is not equal to 1. \
**C.** No it is not a valid PDF because not all densities are nonnegative.  \
**D.** No it is not a valid PDF because both the integral over all possible values is not equal to 1 and not all densities are nonnegative. \
**E.** No it is not a valid PDF because some densities are greater than 1. 

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL

# define a function to evaluate the PDF
...

# create an array w with 1000 evenly spaced values over the range of W
w = ...

# use the function you defined to calculate the PDF for every value in the array w
PDF_w_values = ...

# integrate f(w) over the range for which W is defined. Make sure you index the value of the integral.
integral = ...

# Check the conditions for f(w) to be a valid PDF (check each condition separately to answer the question below)
...

# Is f(w) a valid PDF? Choose your answer from the choices above.
q3 = ...
q3
```

```python
grader.check("q3.0")
```

### Air Particles PDF

Next we will look at a real application of continuous random variables.

Particles are a major component of air pollution and are commonly described by their size. Let $Y$ represent the diameter, in micrometers ($\mu m$ : $10^{-6} m$), of a randomly chosen particle. The probability density function of $Y$ is inversely proportional to the volume of the particle; that is:

$$ f_Y(y)=   \left\{
\begin{array}{ll}
      \Large\frac{2}{y^3} & y \geq 1\\
      0 & \text{otherwise}\\
\end{array} 
\right.  $$

<font color='red'>**Question 4.0.**</font> Define a Python function `PDF_y(y)`, which takes as input a scalar number or a numeric sequence (i.e., list, tuple, NumPy array) of numeric values and returns the density(ies) at the input value(s) using the above equation for $f_Y(y)$. (0.25 pts)

Assume that the input value is $y \geq 1$.

```python
# ANSWER CELL

# define a function to evaluate the PDF
...
```

```python
grader.check("q4.0")
```

<font color='red'>**Question 4.1.**</font> One common measure of air quality is the Particulate Matter 2.5 ($PM_{2.5}$), which refers to particles 2.5 $\mu m$ or less in diameter (i.e., $y<2.5$). Based on the PDF of $Y$, what proportion of the contaminating particles are $PM_{2.5}$? Assign your answer to `q4_1`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.75 pts)

*Hint:* This is the same as getting $P(y\leq2.5)$. How do you get probability from a density function? Be careful regarding the range for which $f_Y(y)$ is defined.

```python
# ANSWER CELL

q4_1 = ...
print(f'PM 2.5 is equal to {q4_1:.3f}' if not isinstance(q4_1, type(Ellipsis)) else None)
```

```python
grader.check("q4.1")
```

<font color='red'>**Question 4.2.**</font> What is the probability that the diameter of a selected air particle is exactly equal to 2.5 $\mu m$. Assign your answer to `q4_2`. (0.5 pts)

*Hint:* This is a conceptual question.

```python
# ANSWER CELL

q4_2 = ...
print(f'P(y = 2.5) = {q4_2}')
```

```python
grader.check("q4.2")
```

<font color='red'>**Question 4.3.**</font> Is it more likely for the diameter of a selected air particle to be close to 5 $\mu m$ than to 10 $\mu m$? Assign your answer to `q4_3` as a string. (0.5 pts)

**A.** We can't tell based on the available information  \
**B.** Yes because $P(Y=5) > P(Y=10)$ \
**C.** No because $P(Y=5) < P(Y=10)$ \
**D.** Yes because $f(Y=5) > f(Y=10)$ \
**E.** No because $f(Y=5) < f(Y=10)$ 

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL

# Type any code below to help you with this question
...

q4_3 = ...
q4_3
```

```python
grader.check("q4.3")
```

### Calculating Numerical Summaries 

In the lecture, we defined the mean of a continuous random variable as:

$\mu_x=E[x]= {\displaystyle \int \limits_{-\infty} ^{+\infty} x \cdot f_X(x) \ dx}\$

When integrating the PDF, we have to be careful of the range over which the PDF is defined. So, if for example, we have:

$$ f_X(x)=   \left\{
\begin{array}{ll}
      1.25(1-x^4) & 0\leq x \leq 1\\
      0 & \text{otherwise}\\
\end{array} 
\right.  $$

Then, to get the mean:

$$E(X) = {\displaystyle \int \limits_{-\infty} ^{+\infty} x \cdot f_X(x) \ dx} =  {\displaystyle \color {red} {\displaystyle \int \limits_{-\infty} ^{0} x \cdot 0 \ dx}} + {\displaystyle \color {blue} {\displaystyle \int \limits_{0} ^{1} x \cdot 1.25(1-x^4) \ dx}} + {\displaystyle \color {red} {\displaystyle \int \limits_{1} ^{+\infty} x \cdot 0 \ dx}} = {\displaystyle \color {blue} {\displaystyle \int \limits_{0} ^{1} x \cdot 1.25(1-x^4) \ dx}}$$

So, we should only integrate $x \cdot 1.25(1-x^4)$ between $0$ and $1$, because this is the range of $X$ for which the PDF equation is defined.

The same applies when integrating the PDF for any other purpose (getting variance, etc.).

Recall that we are working with the following PDF of the diameter, in micrometers ($\mu m$ : $10^{-6} m$), of a randomly chosen air particle:

$$ f_Y(y)=   \left\{
\begin{array}{ll}
      \Large\frac{2}{y^3} & y \geq 1\\
      0 & \text{otherwise}\\
\end{array} 
\right.  $$

You already defined a Python function named `PDF_y()` to get $f_Y(y)$. So to get the mean value, we need to integrate the function $y \cdot f_Y(y)$ over the range of interest. 

To do so, we need to first define a new Python function, say `fun_mean()`, that returns: $y \cdot f_Y(y)$. Then, we can integrate the new function, `fun_mean()` using `integrate.quad()` to compute $\mu_Y$.

<font color='red'>**Question 5.0.**</font> Define a Python function `fun_mean(y)`, which takes as input a scalar number or a numeric sequence (i.e., list, tuple, NumPy array) of numeric values and returns $y \cdot f_Y(y)$ at the input value(s) using the above equation for $f_Y(y)$. (0.25 pts)

Assume that the input value is $y \geq 1$.

```python
# ANSWER CELL

# define a function to evaluate y * f(y)
...
```

```python
grader.check("q5.0")
```

<font color='red'>**Question 5.1.**</font> What is the value of $\mu_Y$? Assign your answer to `mean_y`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.75 pts)

*You have to use `integrate.quad(function_mean, a, b)[0]` to calculate the mean by simply integrating `fun_mean()` (instead of integrating `PDF_y`) over the possible values of $y$. Because $y\geq1$, set the upper limit for the integral equal to $10^6$ instead of $+\infty$.*

```python
# ANSWER CELL

# Use integrate.quad() to calculate the mean
mean_y = ...
print(f'The mean particle diameter is {mean_y:.1f} micrometers' if not isinstance(mean_y, type(Ellipsis)) else None)
```

```python
grader.check("q5.1")
```

We also defined the variance of a continuous random variable as:

$\sigma_X^2=Var[X]= {\displaystyle \int \limits_{-\infty} ^{+\infty} (x-\mu_X)^2 \cdot f_X(x) \ dx }$

In this case, to use `integrate.quad()` to calculate variance, we have to define a function that returns: $(x-\mu_X)^2 \cdot f_X(x)$ and then integrate it.

<font color='red'>**Question 5.2.**</font> Define a Python function `fun_var(y)`, which takes as input a scalar number or a numeric sequence (i.e., list, tuple, NumPy array) of numeric values and returns $(y - \mu_Y)^2 \cdot f_Y(y)$ at the input value(s) using the above equation for $f_Y(y)$. (0.25 pts)

Assume that the input value is $y \geq 1$.

Note that $\mu_Y$ is saved as variable `mean_y`.

```python
# ANSWER CELL

# define a function to evaluate (y - mean)^2 * f(y)
...
```

```python
grader.check("q5.2")
```

<font color='red'>**Question 5.3.**</font> What is the value of $\sigma_Y$? Assign your answer to `stdev_y`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.75 pts)

*Similar to the mean, set the upper limit for the integral equal to $10^6$ instead of $+\infty$.*

```python
# ANSWER CELL

stdev_y = ...
print(f'The standard deviation of the particle diameter is {stdev_y:.3f} micrometers' if not isinstance(stdev_y, type(Ellipsis)) else None)
```

```python
grader.check("q5.3")
```

### Plotting the PDF

Finally, let's plot the probability density function to visually assess the skewness of the distribution of the particle diameter.

We can simply plot the PDF by defining a reasonable range for the $y$ values (using for example, `np.linspace(start, stop, num)`), getting the PDF values, and then plotting them using `plt.plot()`.

<font color='red'>**Question 6.0.**</font> Modify the code below to plot the PDF of $Y$. Follow these steps: (0.75 pts)

1. Create an array named `y` with 1000 evenly spaced values between 1 and 10
2. Calculate the $f_Y(y)$ for every value in the array `y`. Assign the output to `PDF_y_values`. *Hint: use your function PDF_y()*
3. Plot the PDF using `plt.plot(y, PDF_y_values)`. 
4. Change the line color of `plt.plot(y, PDF_y_values)` to anything other than the default color. You can control the line color, style, and other properties. See [here](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.plot.html) for more details.
5. Set the x-axis label to "Particle Diameter (micrometers)"
6. Set the y-axis label to "Density"

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt
from matplotlib.lines import Line2D

# Edit the code below to plot the PDF of Y (only edit where you have ...)

# create an array y with 1000 evenly spaced values between 1 and 10
y = ...

# use the function you defined to calculate the PDF for every value in the array y
PDF_y_values = ...

# create figure and axes
fig_1, ax_1 = plt.subplots(nrows=1, ncols=1, figsize=(5,3))

# Plot the PDF
...

# set title, xlabel, and ylabel
...

# add grid lines
ax_1.grid()

# display all figures
plt.tight_layout()
plt.show()
```

```python
grader.check("q6.0")
```

<font color='red'>**Question 6.1.**</font> Based on your plot of the PDF, what can you say about the skewness of the particle diameter, $Y$? Assign your answer to `q6_1` as a string. (0.25 pts)

**A.** Symmetric\
**B.** Negatively Skewed\
**C.** Positively Skewed

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL

q6_1 = ...
q6_1
```

```python
grader.check("q6.1")
```

### Common Discrete Distributions

After understanding the distribution of the particle diameter, you want to collect several air samples and test them for pollution. In the next few questions, you will be asked to identify the distribution of the random variable, its parameters, and any summary statistics. 




| Choice | Distribution      | Python function |
|:-------|:------------------|:----------------|
| "A"    | Bernoulli         | [`bernoulli(p)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.bernoulli.html#scipy.stats.bernoulli) |
| "B"    | Binomial          | [`binom(n, p)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.binom.html#scipy.stats.binom) |
| "C"    | Hypergeometric    | [`hypergeom(M, n, N)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.hypergeom.html#scipy.stats.hypergeom) |
| "D"    | Geometric         | [`geom(p)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.geom.html#scipy.stats.geom) |
| "E"    | Negative Binomial | [`nbinom(n, p)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.nbinom.html#scipy.stats.nbinom) |
| "F"    | Poisson           | [`poisson(mu)`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.poisson.html#scipy.stats.poisson) |

You collected 10 air samples. Assume that the probability that an air sample has **healthy** levels of $PM_{2.5}$ is the same for each sample and is equal to **0.70**. Also, assume that $PM_{2.5}$ levels in different samples are independent of one another.

Let $X$ represent the number of samples out of the 10 with **unhealthy** levels of $PM_{2.5}$.

<font color='red'>**Question 7.0.**</font> What is the distribution of $X$? Assign your answer to `q7_0` as a string. The letter choices for the different distributions are shown in the table above. (0.25 pts)

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL

q7_0 = ...
q7_0
```

```python
grader.check("q7.0")
```

Python provides specific functions to directly calculate summary statistics for various distributions. 
1. For instance, if you wish to find the mean of a Bernoulli distribution with parameter $p = 0.2$, you can utilize `bernoulli.mean(p=0.2)`.
2. Similarly, to find the median of a Bernoulli distribution with parameter $p = 0.2$, you can utilize `bernoulli.median(p=0.2)`.
3. To obtain the standard deviation of the same Bernoulli distribution with parameter $p = 0.2$, you can use `bernoulli.std(p=0.2)`. 

Make sure to adjust the function names and input parameters accordingly for other distributions. Refer to the table above for the name and parameter(s) of each distribution.

<font color='red'>**Question 7.1.**</font> Calculate the following numerical summaries using appropriate Python functions: (0.75 pts)
1. What is the mean number of samples out of the 10 with unhealthy levels of $PM_{2.5}$, $\mu_X$. Assign the value to `mean_X`.
2. What is the median of the number of samples out of the 10 with unhealthy levels of $PM_{2.5}$. Assign the value to `median_X`.
3. What is the standard deviation of the number of samples out of the 10 with unhealthy levels of $PM_{2.5}$, $\sigma_X$. Assign the value to `stdev_X`.

Do not just manually type the numeric answer. Use Python expressions for the distribution of $X$ that return the desired answer and assign the expressions to the variables.

```python
# ANSWER CELL

mean_X = ...
print(f'The mean of X is {mean_X:.1f} samples' if not isinstance(mean_X, type(Ellipsis)) else None)

median_X = ...
print(f'The median of X is {median_X:.1f} samples' if not isinstance(median_X, type(Ellipsis)) else None)

stdev_X = ...
print(f'The standard deviation of X is {stdev_X:.3f} samples' if not isinstance(stdev_X, type(Ellipsis)) else None)
```

```python
grader.check("q7.1")
```

Now assume that after testing the 10 air samples, 3 were found to have unhealthy levels of $PM_{2.5}$. You forgot to label the samples, and they got mixed up. You selected 5 out of the 10 samples at random and tested them again. Let $V$ be the number of samples with unhealthy levels of $PM_{2.5}$ out of the 5 you tested again. Assume that you selected samples without replacement.

<font color='red'>**Question 8.0.**</font> What is the distribution of $V$? Assign your answer to `q8_0` as a string. The letter choices for the different distributions are shown in the table above. (0.25 pts)

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL

q8_0 = ...
q8_0
```

```python
grader.check("q8.0")
```

<font color='red'>**Question 8.1.**</font> Calculate the following numerical summaries using appropriate Python functions: (0.75 pts)
1. What is the mean of the number of samples with unhealthy $PM_{2.5}$ levels out of the 5 you tested again, $\mu_V$. Assign the value to `mean_V`.
2. What is the median of the number of samples with unhealthy $PM_{2.5}$ levels out of the 5 you tested again,. Assign the value to `median_V`.
3. What is the standard deviation of the number of samples with unhealthy $PM_{2.5}$ levels out of the 5 you tested again,, $\sigma_V$. Assign the value to `stdev_V`.


**Note that the input parameters might use different symbols or be different from the ones we defined in the lecture.** So it's important you read the corresponding documentation carefully by clicking on the Python function of the corresponding distribution in the table above.

```python
# ANSWER CELL

mean_V = ...
print(f'The mean of V is {mean_V:.1f} samples' if not isinstance(mean_V, type(Ellipsis)) else None)

median_V = ...
print(f'The median of V is {median_V:.1f} samples' if not isinstance(median_V, type(Ellipsis)) else None)

stdev_V = ...
print(f'The standard deviation of V is {stdev_V:.3f} samples' if not isinstance(stdev_V, type(Ellipsis)) else None)
```

```python
grader.check("q8.1")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Finally, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing this lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 0.5)
```
>>> assert isinstance(q1_0, tuple)
>>> assert all((isinstance(option, str) for option in q1_0))
>>> assert all((len(option) == 1 for option in q1_0))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('800618943025315f869e4e1f09471012' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.


Tests for q2.0 (points: 0.5)
```
>>> assert get_hash(PDF_z(1)) == '46aeed2a0e482947e98b93752353c7cc'
>>> assert get_hash(PDF_z(5)) == 'cb5ae17636e975f9bf71ddf5bc542075'
>>> assert get_hash(PDF_z(10)) == 'c7e230ff5bea174a9dad0a5f34a3d48f'
```
failure_message: Incorrect function.
```
>>> assert get_hash(type(q2_0)) != '14e736438b115821cbb9b7ac0ba79034'
```
failure_message: Make sure to test your function!
```
>>> assert get_hash(q2_0) == '29d679c60745d297cf5863ac1ce9ec74'
```
failure_message: Incorrect q2_0.


Tests for q2.1 (points: 0.5)
```
>>> assert get_hash(len(z)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(z), 15)) == 'aa772fc56cc61aadc9e536e936619f89'
```
failure_message: Incorrect z.
```
>>> assert get_hash(len(PDF_z_values)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(PDF_z_values), 15)) == '62dca49f0781bf26b4305bddb0414bea'
```
failure_message: Incorrect PDF_z_values. Simply call your function PDF_z() uisng z as input.


Tests for q2.2 (points: 0.5)
```
>>> assert get_hash(round(P_0_3, 15)) == '8f11bfb91ec29936603314c7cbc46119'
```
failure_message: Incorrect P_0_3.
```
>>> assert get_hash(round(P_5_9, 15)) == '6ca65d752da20bfe10fa409a89b0379b'
```
failure_message: Incorrect P_5_9.


Tests for q3.0 (points: 1.0)
```
>>> assert get_hash(PDF_w(0)) == '4c8d9ac7c62b7fbf21fb2f2ae14dd65b'
>>> assert get_hash(PDF_w(0.5)) == 'b247775f8ab3d659c1aa9bf453cd0912'
>>> assert get_hash(PDF_w(1)) == '30565a8911a6bb487e3745c0ea3c8224'
```
failure_message: Incorrect function.
```
>>> assert get_hash(len(w)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(w), 15)) == '7ee66a0b059261539c7a25b7cfa5012b'
```
failure_message: Incorrect w.
```
>>> assert get_hash(len(PDF_w_values)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(PDF_w_values), 15)) == 'c56d654ac7d1a88aa93f1a90787e46e0'
```
failure_message: Incorrect PDF_w_values. Simply call your function PDF_w() uisng w as input.
```
>>> assert get_hash(round(integral, 15)) == 'd310cb367d993fb6fb584b198a2fd72c'
```
failure_message: Incorrect integral. You should integrate f(w) over its range and only index the value of the integral.
```
>>> assert len(q3) == 1
>>> assert isinstance(q3, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q3.upper()) == '9d5ed678fe57bcca610140957afab571'
>>> assert get_hash(round(sum(PDF_w_values), 15)) == 'c56d654ac7d1a88aa93f1a90787e46e0'
>>> assert get_hash(round(integral, 15)) == 'd310cb367d993fb6fb584b198a2fd72c'
```
failure_message: Incorrect answer.


Tests for q4.0 (points: 0.25)
```
>>> assert get_hash(PDF_y(1)) == 'd1bd83a33f1a841ab7fda32449746cc4'
>>> assert get_hash(PDF_y(5)) == 'ba4b140343c1a0b1ad0dd842f6b6a4f6'
>>> assert get_hash(PDF_y(10)) == 'd1c6c5ed1452a2c20690b2d93a4f00e0'
```
failure_message: Incorrect function.


Tests for q4.1 (points: 0.75)
```
>>> assert get_hash(round(q4_1, 10)) == '8b302979fdd1a5811d5e4f96d0f6dccd'
```
failure_message: Incorrect answer.


Tests for q4.2 (points: 0.5)
```
>>> assert get_hash(type(q4_2)) != '781f5f2052cfb496fa73d70ea1fbdaa8'
>>> assert get_hash(round(q4_2, 10)) == 'cfcd208495d565ef66e7dff9f98764da'
```
failure_message: Incorrect answer.


Tests for q4.3 (points: 0.5)
```
>>> assert len(q4_3) == 1
>>> assert isinstance(q4_3, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q4_3.upper()) == 'f623e75af30e62bbd73d6df5b50bb7b5'
>>> assert get_hash(PDF_y(1)) == 'd1bd83a33f1a841ab7fda32449746cc4'
>>> assert get_hash(PDF_y(5)) == 'ba4b140343c1a0b1ad0dd842f6b6a4f6'
>>> assert get_hash(PDF_y(10)) == 'd1c6c5ed1452a2c20690b2d93a4f00e0'
```
failure_message: Incorrect answer.


Tests for q5.0 (points: 0.25)
```
>>> assert get_hash(fun_mean(1)) == 'd1bd83a33f1a841ab7fda32449746cc4'
>>> assert get_hash(fun_mean(5)) == '37d5c0b6b354bc3c790d2696b42756c9'
>>> assert get_hash(fun_mean(10)) == '0d16a6cb2ca34d65da949608a7bc01d1'
```
failure_message: Incorrect function.


Tests for q5.1 (points: 0.75)
```
>>> assert get_hash(round(mean_y, 10)) == '8569c22c1bc7a8fbd48c31cb624b9479'
```
failure_message: Incorrect answer. Make sure you are integrating fun_mean over the correct limits.


Tests for q5.2 (points: 0.25)
```
>>> assert get_hash(fun_var(1)) == '9833e8443b6ce951636614be724b0151'
>>> assert get_hash(fun_var(5)) == '8f6e982d6e7139006039aed9bc2d0768'
>>> assert get_hash(fun_var(10)) == '3cccf2a7236332d329f98495973ba7bd'
```
failure_message: Incorrect function.


Tests for q5.3 (points: 0.75)
```
>>> assert get_hash(round(stdev_y, 10)) == '488ab874694a72822348d6cf155fb487'
```
failure_message: Incorrect answer. Make sure you are integrating fun_var over the correct limits. Also, you need standard deviation.


Tests for q6.0 (points: 0.75)
```
>>> assert get_hash(len(y)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(y), 15)) == '69199d4fe1973959c323b0d5450baf1f'
```
failure_message: Incorrect y.
```
>>> assert get_hash(len(PDF_y_values)) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(round(sum(PDF_y_values), 15)) == '3a8a83ce01f872534e3cabd4ce147e33'
```
failure_message: Incorrect PDF_y_values. Simply call your function PDF_y() uisng y as input.
```
>>> assert type(fig_1).__name__ == 'Figure'
>>> assert type(ax_1).__name__ == 'Axes'
```
failure_message: Make sure to create a figure with 1 axes and assign them to "fig_1" and "ax_1".
```
>>> assert isinstance(ax_1.lines[0], Line2D)
```
failure_message: Make sure you are plotting to ax_1 using plot().
```
>>> assert get_hash(sum(ax_1.lines[0].get_xdata())) == '69199d4fe1973959c323b0d5450baf1f'
>>> assert get_hash(len(ax_1.lines[0].get_xdata())) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(sum(ax_1.lines[0].get_ydata())) == '3a8a83ce01f872534e3cabd4ce147e33'
>>> assert get_hash(len(ax_1.lines[0].get_ydata())) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
```
failure_message: Check the arguments of plt.plot()
```
>>> assert ax_1.lines[0].get_color() != '#1f77b4'
>>> assert ax_1.lines[0].get_linestyle() == '-'
```
failure_message: Change the line color inside plot() to any color other than the default. Do not change line style.
```
>>> keywords = ['PARTICLE', 'DIAMETER', 'MICROMETERS']
>>> assert all((word in ax_1.get_xlabel().upper() for word in keywords))
```
failure_message: Check x-axis label. Make sure correct units are included.
```
>>> keywords = ['DENSITY']
>>> assert all((word in ax_1.get_ylabel().upper() for word in keywords))
```
failure_message: Check y-axis label.


Tests for q6.1 (points: 0.25)
```
>>> assert len(q6_1) == 1
>>> assert isinstance(q6_1, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q6_1.upper()) == '0d61f8370cad1d412f80b84d143e1257'
>>> assert get_hash(sum(ax_1.lines[0].get_xdata())) == '69199d4fe1973959c323b0d5450baf1f'
>>> assert get_hash(len(ax_1.lines[0].get_xdata())) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
>>> assert get_hash(sum(ax_1.lines[0].get_ydata())) == '3a8a83ce01f872534e3cabd4ce147e33'
>>> assert get_hash(len(ax_1.lines[0].get_ydata())) == 'a9b7ba70783b617e9998dc4dd82eb3c5'
```
failure_message: Incorrect answer.


Tests for q7.0 (points: 0.25)
```
>>> assert len(q7_0) == 1
>>> assert isinstance(q7_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q7_0.upper()) == '9d5ed678fe57bcca610140957afab571'
```
failure_message: Incorrect answer.


Tests for q7.1 (points: 0.75)
```
>>> assert get_hash(round(mean_X, 10)) == '55c82b601deae028c1c5e87fd820923d'
```
failure_message: Incorrect mean.
```
>>> assert get_hash(round(median_X, 10)) == '55c82b601deae028c1c5e87fd820923d'
```
failure_message: Incorrect median.
```
>>> assert get_hash(round(stdev_X, 10)) == '383951d96f2cba387df6ba40e94f885d'
```
failure_message: Incorrect standard deviation.


Tests for q8.0 (points: 0.25)
```
>>> assert len(q8_0) == 1
>>> assert isinstance(q8_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q8_0.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: Incorrect answer.


Tests for q8.1 (points: 0.75)
```
>>> assert get_hash(round(mean_V, 10)) == '6008647277c4454cecd97d33c069f0ca'
```
failure_message: Incorrect mean.
```
>>> assert get_hash(round(median_V, 10)) == 'e4c2e8edac362acab7123654b9e73432'
```
failure_message: Incorrect median.
```
>>> assert get_hash(round(stdev_V, 10)) == '663f9fb4061bfa1081f26c0282f70d19'
```
failure_message: Incorrect standard deviation.
