```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# CE 93: Lab Assignment 04

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to learn about properties of discrete random variables. We will not load any data. Instead, we will create the random variables ourselves.

## Instructions 
**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and these should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

import math
import numpy as np
import pandas as pd
import statistics as stats
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import resources
```

### Properties of Probability Mass Function 

Let's first review some basic concepts.

<font color='red'>**Question 1.0.**</font> What is(are) the property(ies) of a probability mass function? 
 Assign ALL that apply to the variable `q1_0`. (0.5 pts)

**A**. The probability mass function is a function that assigns probabilities to each possible outcome of a discrete random variable \
**B**. The probability of any outcome can be greater than 1 as long as the probability mass function adds up to 1 \
**C**. The probability mass function of any discrete random variable should add up to 1 \
**D**. The probability of any outcome should be greater than or equal to 0 and less than or equal to 1 \
**E**. The probability of any outcome can be negative as long as the probability mass function adds up to 1 \
**F**. The probability of any outcome should be strictly greater than 0 and strictly less than 1

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q1_0 = ...
q1_0
```

```python
grader.check("q1.0")
```

### Random Variable Example

You have been asked to investigate the daily temperature trends in Berkeley. Specifically, you want to study the number of days that will have an average temperature above normal for a specific week in February. Let $X$ be the **number of days** that have an average temperature above normal for a week in February.

<font color='red'>**Question 1.1.**</font> Is $X$ a discrete or continuous random variable? Assign the boolean value `True` to `q1_1` if it is discrete and `False` if it is continuous. (0.25 pts)

```python
# ANSWER CELL
q1_1 = ...
q1_1
```

```python
grader.check("q1.1")
```

### Probability Mass Function

Based on historical data for February, you were able to calculate the probability mass function (PMF) for $X$. The PMF is shown in the table below. After being done with your calculations, you realized that you forgot to calculate the probability for $X=3$.


|$X$|0|1|2|3|4|5|6|7|
|:-:|:-:|:-:|:-:|:-:|:-:|:-:|:-:|:-:|
|PMF|0.05|0.19|0.15|??|0.15|0.20|0.008|0.005|

<font color='red'>**Question 2.0.**</font> Using properties of a PMF, which of the following options can be used to calculate $p(X=3)$ in `Python`? Assign your answer to `q2_0` using exactly one of the expressions below as it is. Do not just manually type the numeric answer. Use the Python expression that you select and assign it to the variable. Do **not** round your answer. (0.5 pts)

* `1 + 0.05 + 0.19 + 0.15 + 0.15 + 0.2 + 0.008 + 0.005`
* `1 - 0.05 - 0.19 - 0.15 - 0.15 - 0.2 - 0.008 - 0.005`
* `0.05 + 0.19 + 0.15 + 0.15 + 0.2 + 0.008 + 0.005`
* `1 - 0.15`
* `1 - 0.15 - 0.15`

```python
# ANSWER CELL
q2_0 = ...
print(f'P(X = 3) = {q2_0}')
```

```python
grader.check("q2.0")
```

<font color='red'>**Question 2.1.**</font> Calculate $p(X=3)$ by hand or by using a calculator. For this problem, please manually type your numeric answer and assign it to `q2_1`. (0.25 pts)

```python
# ANSWER CELL
q2_1 = ...
print(f'P(X = 3) = {q2_1}')
```

```python
grader.check("q2.1")
```

### Why Are My Values Slightly Different ?!

You should see a very slight error in the `Python` result. This is know as a **round-off error** and is due to the fact that floating point numbers (i.e., a number with a decimal point) cannot always be stored with perfect precision in computers. 

Let's try a simpler example. We know that $1.0-0.95$ is equal to $0.05$. Let's test that in `Python`. To check if two expressions are equal in `Python`, we can use `==` (double equal sign). 

Copy and paste this line into the cell below and run it `1.0 - 0.95 == 0.05`. 
Running `1.0 - 0.95 == 0.05` is equivalent to asking , "*Is $1.0-0.95$ equal to  $0.05$?*" Since this statement is true, Python, ideally, should return the value `True`.

So why does `Python` return `False` in this case? We know that the above statement should be `True`. Read [this short page](https://pythonnumericalmethods.berkeley.edu/notebooks/chapter09.03-Roundoff-Errors.html) to understand round-off error and to answer the next question.

```python
# Try the above example

```

<font color='red'>**Question 2.2.**</font> Which of the following commands give us the exact value of $p(X=3)$ (i.e., the value in Question 2.1.) using `Python`? Assign ALL that apply to the variable `q2_2`. (0.5 pts)

Note that when you calculated $p(X=3)$ in `Python`, you assigned the answer to the variable `q2_0`. 

**A.** `round(q2_0)` \
**B.** `round(q2_0, 2)` \
**C.** `round(q2_0, 3)` \
**D.** `round(q2_0, 5)` 

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# Type any code below to help you with this question.

```

```python
# ANSWER CELL
q2_2 = ...
q2_2
```

```python
grader.check("q2.2")
```

### Create Numpy Arrays

Recall that the PMF of $X$ (with the missing probability) is:

|$X$|0|1|2|3|4|5|6|7|
|:-:|:-:|:-:|:-:|:-:|:-:|:-:|:-:|:-:|
|PMF|0.05|0.19|0.15|??|0.15|0.20|0.008|0.005|

We want to create a numpy array with the PMF values of $X$ using `np.array()`. For example, if we want to create a numpy array with the values 0.25, 0.50, 0.25, we would use: `np.array([0.25, 0.50, 0.25])`. So, we simply list the PMF values between square brackets and separate them with commas.

<font color='red'>**Question 3.0.**</font> Create a numpy array with the PMF values of $X$ using `np.array()`. Assign your answer to `PMF_X`. (0.5 pts)

When creating the numpy array, make sure that the PMF values are **in order**. So the first value is for $X=0$, then $X=1$, etc. Don't forget to **include** the correct value for $p(X=3)$, which you calculated above (Question 2.1.).

```python
# ANSWER CELL
PMF_X = ...
print(f'PMF of X: {PMF_X}')
```

```python
grader.check("q3.0")
```

### Example PMF Plot

We can plot a PMF in Python using `plt.vlines()`. Run the code cell below to look at an example for a simplified PMF with probabilities equal to `[0.25, 0.50, 0.25]`. You can read more about `plt.vlines(x, ymin, ymax)` [here](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.vlines.html). As the name suggests, it plots vertical lines at each `x` starting from `ymin` to `ymax`.

```python
# Example PMF plot

# create a numpy array of the PMF
PMF_example = np.array([0.25, 0.50, 0.25])

# Assume that the PMF above is for z = 0, z = 1, and z = 2. 
# We can create a new numpy array for the Z values, or simply use np.arange().
# np.arange(end) will create a numpy array with values [0, 1, ..., end-1]
Z = np.arange(3)

# create figure and axes
fig_1, ax_1 = plt.subplots(nrows=1, ncols=1, figsize=(3,2.5))

# Plot PMF using plt.vlines(x, ymin, ymax)
# Specify line color = blue, line style = solid line, and line width = 3
plt.vlines(Z, 0, PMF_example, colors='b', linestyles='-', lw=3)

# set title, xlabel, ylabel, y-axis limits, and locations of x-axis tick marks
ax_1.set(title='PMF Example', 
         xlabel='$Z$', 
         ylabel='Probability',
         ylim=[0, 0.6],
         xticks=Z)

# display plot
plt.show()
```

### PMF Plot of $X$

<font color='red'>**Question 4.0.**</font> Modify the code below to plot the PMF of the random variable $X$ (number of days that have an average temperature above normal for a week in February). Recall that the PMF of $X$ is saved as `PMF_X`. (1 pt)

1. Set the x-axis label to "X".
2. Set the y-axis label to "Probability".
3. Set the title to "PMF of X".

Specify appropriate limits for both axes.

You can refer to the code above for an example.

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt
from matplotlib.collections import LineCollection

# create figure and axes
fig_2, ax_2 = plt.subplots(nrows=1, ncols=1, figsize=(5,2.5))

# define values for X
X = ...

# insert your code to plot the PMF of X
# set xlabel, ylabel, title, and any other properties
...

...

# display the figure 
plt.show()
```

```python
grader.check("q4.0")
```

### Cumulative Distribution Function

We want to compute the cumulative distribution function (CDF) of $X$. Let's first review the basic properties of a CDF.

<font color='red'>**Question 5.0.**</font> What is(are) the property(ies) of a cumulative distribution function? 
 Assign ALL that apply to the variable `q5_0`. (0.5 pts)

**A**. The cumulative distribution function is a function that gives the probability that a discrete random variable is less than or equal to a specific value \
**B**. The cumulative distribution function of any discrete random variable should add up to 1 \
**C**. The values of the cumulative distribution function should always be greater than or equal to 0 and less than or equal to 1 \
**D**. For all discrete random variables, $F(-\infty)=0$  and $F(+\infty)=1$ \
**E**. The cumulative distribution function may be increasing, decreasing, or constant, as long as it reaches a final value of 1\
**F**. The cumulative distribution function is a non-decreasing function

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.er choice.

```python
# ANSWER CELL
q5_0 = ...
q5_0
```

```python
grader.check("q5.0")
```

If we have the PMF, then the CDF can be calculated as follows:

$$ F(x)= P(X\leq x)= \sum \limits _{x_i<x} ^{} p(x_i)$$

`Python` provides several easy functions to calculate the cumulative sum.

<font color='red'>**Question 5.1.**</font> Which of the following commands can be used to correctly calculate the CDF of $X$? Assign ALL that apply to the variable `q5_1`. Assume that we only want the CDF at the given values of $X$ (i.e., 0, 1, 2, ..., 7). (0.5 pts)

**A.** `np.cumsum(PMF_X)` \
**B.** `sum(PMF_X)` \
**C.** `PMF_X.cumsum()`

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

*You can try out the different functions and compare their output with a hand calculation.*

```python
# Try out different ways to compute CDF

```

```python
# ANSWER CELL
q5_1 = ...
q5_1
```

```python
grader.check("q5.1")
```

<font color='red'>**Question 5.2.**</font> Compute the CDF of $X$ using an appropriate function from the options above. Assign your answer to `CDF_X`. Do not just manually type the numeric answers. Use Python expressions that return the desired answer and assign the expression to the variable. (0.5 pts)

```python
# ANSWER CELL
# Calculate the CDF of X
CDF_X = ...
print(f'CDF of X: {CDF_X}')
```

```python
grader.check("q5.2")
```

###  CDF Plot of $X$ 

Once we have computed the CDF values, we can plot it in `Python` using `plt.step()`.

<font color='red'>**Question 6.0.**</font> Modify the code below to plot the CDF of $X$ using the following command: `plt.step(X, CDF_X, c='r')`

Is the plot correct? Assign the boolean value `True` to the variable `q6_0` if the plot is correct and `False` if it is incorrect. (0.5 pts)

*Hint: Refer to the lecture notes to see the relation between CDF and PMF. The CDF should show "jumps" at each possible value of the random variable, and the value of the jump should be equal to the value of the PMF at the corresponding $x$ value. So, at $x_i$, the CDF should show a jump and the value of the jump should be equal to $p(x_i)$. In the same plot of the CDF, also plot the PMF using `plt.vlines(X, 0, PMF_X)`. Then, check to see where the jumps of the CDF occur and if their values are equal to the corresponding $p(x_i)$.*

*Note that `X` in `plt.step(X, CDF_X, c='r')` is simply an array with all $X$ values in order (i.e., 0, 1, 2, etc.). You should have already created this array to plot the PMF.*

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt
from matplotlib.lines import Line2D

# create figure and axes
fig_3, ax_3 = plt.subplots(nrows=1, ncols=1, figsize=(5,3))

# insert your code to plot the CDF of X
# also plot the PMF of X for easier comparison
# set xlabel, ylabel, title, and any other properties
...
...

...

# display the figure 
plt.show()
```

```python
# ANSWER CELL
q6_0 = ...
q6_0
```

```python
grader.check("q6.0")
```

Well, would we ask you if the plot is correct if it were actually correct?! Now, let us (aka you) figure out how to plot it correctly. Read about the parameters for `plt.step()` [here](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.step.html). Specifically, read and understand how the parameter `where` works.

<font color='red'>**Question 6.1.**</font> Modify the code below to **correctly** plot the CDF of $X$. (0.5 pts)

*To specify input parameters for a function, you simply add them between parentheses. For example, in `plt.step(X, CDF_X, c='r')`, we added `c='r'`, which is an input parameter that specifies the color of the line as red. Likewise, you can specify the input parameter `where` by adding it between parentheses after a comma.* 

*Try out the different options for `where` and then check to see which option shows the correct jumps of the CDF. You can also include in the same plot the PMF to make it easier to compare the values of the jumps.*

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt
from matplotlib.lines import Line2D

# create figure and axes
fig_4, ax_4 = plt.subplots(nrows=1, ncols=1, figsize=(5,3))

# insert your code to plot the CDF of X
# also plot the PMF of X for easier comparison
# set xlabel, ylabel, title, and any other properties
...
...

...

# display the figure 
plt.show()
```

```python
grader.check("q6.1")
```

### Extra: Better Looking CDF
We can plot the CDF in a more polished manner. Run the code below to check the output. Feel free to explore the code on your own or ask us any questions you may have! You do not have to answer any questions for this plot.

```python
# Run this cell to plot CDF of X

# create figure and axes
fig_5, ax_5 = plt.subplots(nrows=1, ncols=1, figsize=(5,3))

# append values to X to plot lines extending below 0 and greater than the maximum
X_appended = np.arange(-1,9,1)

# insert a value 0 at the the first position (i.e., 0) of CDF_X
CDF_X_appended = np.insert(CDF_X, 0, 0)

# plot horizontal lines for CDF
plt.hlines(y=CDF_X_appended, xmin=X_appended[:-1], xmax=X_appended[1:], color='r', zorder=1)

# Plot white filled points
plt.scatter(X_appended[1:-1], CDF_X_appended[:-1], color='white', s=20, zorder=2, edgecolor='r')

# Plot solid filled points
plt.scatter(X_appended[1:-1], CDF_X, color='r', s=20, zorder=2)

# set title, axes lables, axes limits
ax_5.set(title='Better Looking CDF',
         xlabel='$X$', 
         ylabel='CDF',
         ylim=[-0.05, 1.05],
         xlim=[-.9, 7.9])

# add grid lines on the y-axis
plt.grid(axis="y")

# display the figure 
plt.show()
```

### Probabilities

Next, you will be asked to obtain the probabilities of different events based on the PMF and CDF of the number of days that have an average temperature above normal for a week in February, $X$.

<font color='red'>**Question 7.0.**</font> What is the probability that at most two days will have an average temperature above normal for a week in February, $P(X\leq 2)$? Assign your numeric answer to `q7_0`. You shouldn't be doing any new calculations. Refer to the PMF and CDF of $X$ and then manually type the numeric answer. (0.25 pts)

```python
# ANSWER CELL
q7_0 = ...
f'The probability that at most two days will have an average temperature above normal is {q7_0:.3f}' if not isinstance(q7_0, type(Ellipsis)) else None
```

```python
grader.check("q7.0")
```

<font color='red'>**Question 7.1.**</font> What is the probability that less than two days will have an average temperature above normal for a week in February, $P(X < 2)$? Assign your numeric answer to `q7_1`. You shouldn't be doing any new calculations. Refer to the PMF and CDF of $X$ and then manually type the numeric answer. (0.25 pts)

```python
# ANSWER CELL
q7_1 = ...
f'The probability that less than two days will have an average temperature above normal is {q7_1:.3f}' if not isinstance(q7_1, type(Ellipsis)) else None
```

```python
grader.check("q7.1")
```

<font color='red'>**Question 7.2.**</font> What is the probability that more than two days will have an average temperature above normal for a week in February, $P(X > 2)$? Assign your answer to `q7_2`. Use Python expressions that return the desired answer and assign the expression to the variable.(0.25 pts)

```python
# ANSWER CELL
q7_2 = ...
f'The probability that more than two days will have an average temperature above normal is {q7_2:.3f}' if not isinstance(q7_2, type(Ellipsis)) else None
```

```python
grader.check("q7.2")
```

<font color='red'>**Question 7.3.**</font> What is the probability that less than 3.5 days will have an average temperature above normal for a week in February, $P(X < 3.5)$? Assign your numeric answer to `q7_3`. You shouldn't be doing any new calculations. Refer to the PMF and CDF of $X$ and then manually type the numeric answer. (0.5 pts)

```python
# ANSWER CELL
q7_3 = ...
f'The probability that less than 3.5 days will have an average temperature above normal is {q7_3:.3f}' if not isinstance(q7_3, type(Ellipsis)) else None
```

```python
grader.check("q7.3")
```

### Calculating Numerical Summaries 

In the lecture, we defined the mean of a discrete random variable as:

$\mu_X=E[X]= \sum \limits_{\text{all } x_i}^{} x_i \cdot p(x_i) = x_1 \cdot p(x_1) + ... + x_n \cdot p(x_n)$

To calculate this, each $x_i$ value needs to be multiplied by its corresponding $p(x_i)$, and then the sum of all $x_i \cdot p(x_i)$ is taken.

If you have a numpy array containing all possible values of the random variable, let's call it array `X`, and another numpy array of the corresponding probabilities, let's call it array `PMF_X`, then the expression `X * PMF_X` will perform element-wise multiplication. It will return an array with values $[x_1 \cdot p(x_1), \ x_2 \cdot p(x_2), \ ..., \ x_n \cdot p(x_n)]$. Finally, the mean of $X$ can be computed by summing the values of this new array, which can be done in Python using the `sum()` function.

So, `sum(X * PMF_X)` is essentially equivalent to: $\sum \limits_{\text{all } x_i}^{} x_i \cdot p(x_i)$.

<font color='red'>**Question 8.0.**</font> Calculate the mean value of the number of days that have an average temperature above normal for a week in February, $\mu_X$. Assign your answer to `mean_X`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.5 pts)

```python
# ANSWER CELL
mean_X = ...
f'The mean of X is {mean_X:.3f} days' if not isinstance(mean_X, type(Ellipsis)) else None
```

```python
grader.check("q8.0")
```

The variance of a discrete random variable can be determined using the following formula:

$\sigma_X^2=Var[X]= \sum \limits_{\text{all } x_i}^{} (x_i-\mu_X)^2 \cdot p(x_i) = (x_1-\mu_X)^2 p(x_1) + ... + (x_n-\mu_X)^2 p(x_n)$

To compute the variance in Python, we can utilize operators and functions similar to what we did for the mean.

In this case, the expression `sum((X - mean_X)**2 * PMF_X)` is analogous to the mathematical expression $\sum \limits_{\text{all } x_i}^{} (x_i-\mu_X)^2 \cdot p(x_i)$, where `X` is a numpy array of the values of the random variable, `mean_X` is a scalar that represents the mean of $X$, and `PMF_X` is a numpy array of the corresponding probabilities.

<font color='red'>**Question 8.1.**</font> Calculate the **standard deviation** of the number of days that have an average temperature above normal for a week in February, $\sigma_X$. Assign your answer to `stdev_X`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.5 pts)

```python
# ANSWER CELL
stdev_X = ...
f'The standard deviation of X is {stdev_X:.3f} days' if not isinstance(stdev_X, type(Ellipsis)) else None
```

```python
grader.check("q8.1")
```

Lastly, skewness can be determined using the following formula:

$\theta_X = {\displaystyle\sum \limits_{\text{all } x_i}^{} \left[\left(\frac{x_i-\mu_X}{\sigma_X}\right)^3 \cdot p(x_i)\right]}$

Recall that skewness is a numerical parameter that describes whether the distribution is left/negatively skewed, symmetric, or right/positively skewed.

<font color='red'>**Question 8.2.**</font> Calculate the **skewness** of the number of days that have an average temperature above normal for a week in February, $\theta_X$. Assign your answer to `skew_X`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. (0.5 pts)

*Hint: Use the above equation and modify it to use `sum()`, multiplication, and exponentiation, just like the mean and variance equations.*

```python
# ANSWER CELL
skew_X = ...
f'The skewness of X is {skew_X:.3f}' if not isinstance(skew_X, type(Ellipsis)) else None
```

```python
grader.check("q8.2")
```

It is common to use the following guidelines to describe a distribution based on the skewness coefficient, $\theta_X$:
* $\theta_X < -1$, then $X$ is highly negatively skewed
* $-1 < \theta_X < -0.5$, then $X$ is moderately negatively skewed
* $-0.5 < \theta_X < 0.5$, then $X$ is symmetric
* $0.5< \theta_X < 1$, then $X$ is moderately positively skewed
* $\theta_X > 1$, then $X$ is highly positively skewed

<font color='red'>**Question 8.3.**</font> Based on the skewness value, $\theta_X$, and the above definitions, what can you say about $X$? Assign your answer to the variable `q8_3` as a string. (0.25 pts)

**A.** Symmetric \
**B.** Highly negatively skewed  \
**C.** Moderately negatively skewed  \
**D.** Moderately positively skewed  \
**E.** Highly positively skewed 

Your answer should be a string, e.g., `"A"`, `"B"`, etc.\
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q8_3 = ...
q8_3
```

```python
grader.check("q8.3")
```

Recall that $X$ represents the number of days that have an average temperature above normal for **a week** in February. We are interested in getting summary statistics regarding the number of days that have an average temperature above normal for **the full month** of February. Assume that February has **exactly four weeks**.

<font color='red'>**Question 8.4.**</font> Calculate the **mean** of the number of days that have an average temperature above normal for the full month of February. Assign your answer to `mean_month`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. Recall that the mean for one week is saved as variable `mean_X`. Do not approximate `mean_X`. Instead, use the actual value which is saved as variable `mean_X` in your calculation. (0.5 pts)

```python
# ANSWER CELL
mean_month = ...
f'The mean for the full month is {mean_month:.3f} days' if not isinstance(mean_month, type(Ellipsis)) else None
```

```python
grader.check("q8.4")
```

<font color='red'>**Question 8.5.**</font> Calculate the **standard deviation** of the number of days that have an average temperature above normal for the full month of February. Assign your answer to `stdev_month`. Do not just manually type the numeric answer. Use Python expressions that return the desired answer and assign the expression to the variable. Recall that the standard deviation for one week is saved as variable `stdev_X`. Do not approximate `stdev_X`. Instead, use the actual value which is saved as variable `stdev_X` in your calculation. (0.5 pts)

```python
# ANSWER CELL
stdev_month = ...
f'The standard deviation for the full month is {stdev_month:.3f} days' if not isinstance(stdev_month, type(Ellipsis)) else None
```

```python
grader.check("q8.5")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Finally, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing this lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 0.5)
```
>>> assert isinstance(q1_0, tuple)
>>> assert all((isinstance(option, str) for option in q1_0))
>>> assert all((len(option) == 1 for option in q1_0))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('800618943025315f869e4e1f09471012' == get_hash(option.upper()) for option in q1_0))
```
failure_message: Incorrect answer.


Tests for q1.1 (points: 0.25)
```
>>> assert get_hash(q1_1) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Incorrect answer.


Tests for q2.0 (points: 0.5)
```
>>> assert get_hash(round(q2_0, 16)) == '47323653b0e324130f69269940d8720d'
```
failure_message: Incorrect answer.


Tests for q2.1 (points: 0.25)
```
>>> assert get_hash(q2_1) == '159b75711ed4991efd775e327e6b048c'
```
failure_message: Incorrect answer.


Tests for q2.2 (points: 0.5)
```
>>> assert isinstance(q2_2, tuple)
>>> assert all((isinstance(option, str) for option in q2_2))
>>> assert all((len(option) == 1 for option in q2_2))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q2_2))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q2_2))
```
failure_message: Incorrect answer.
```
>>> assert not any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q2_2))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q2_2))
```
failure_message: Incorrect answer.


Tests for q3.0 (points: 0.5)
```
>>> assert get_hash(len(PMF_X)) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(sum(PMF_X)) == 'e4c2e8edac362acab7123654b9e73432'
>>> assert get_hash(round(PMF_X[0], 3)) == 'b14399cbaac6da4b5b733b483106383f'
>>> assert get_hash(round(PMF_X[1], 3)) == '2f29acc3347fd6f6ff001c1337a94fd2'
>>> assert get_hash(round(PMF_X[2], 3)) == '9ffc7a03c1959674fa8c0a7369460097'
>>> assert get_hash(round(PMF_X[3], 3)) == '159b75711ed4991efd775e327e6b048c'
>>> assert get_hash(round(PMF_X[4], 3)) == '9ffc7a03c1959674fa8c0a7369460097'
>>> assert get_hash(round(PMF_X[5], 3)) == '3d522deaf85577451c01974654b36ad3'
>>> assert get_hash(round(PMF_X[6], 3)) == 'e0d088d01b309e29e7e12e7b657d04c6'
>>> assert get_hash(round(PMF_X[7], 3)) == 'f4881cfc4c6112c01b7cbb980bf0175c'
```
failure_message: Incorrect answer.


Tests for q4.0 (points: 1.0)
```
>>> assert type(fig_2).__name__ == 'Figure'
>>> assert type(ax_2).__name__ == 'Axes'
```
failure_message: Make sure to create a figure with 1 axes and assign them to "fig_2" and "ax_2".
```
>>> assert isinstance(ax_2.collections[0], LineCollection)
```
failure_message: Make sure you are plotting to ax_2 using vlines().
```
>>> assert get_hash(len(X)) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(sum(X)) == '33e75ff09dd601bbe69f351039152189'
```
failure_message: Check X values. Refer to the table above.
```
>>> assert get_hash(len(ax_2.collections[0].get_segments())) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert all((segment.size == 4 for segment in ax_2.collections[0].get_segments()))
>>> assert all((np.unique(segment[:, 0])[0] == count for segment, count in zip(ax_2.collections[0].get_segments(), np.arange(8))))
>>> assert all((np.isclose(np.sum(segment[:, 1]), value, atol=0.001) for segment, value in zip(ax_2.collections[0].get_segments(), PMF_X)))
```
failure_message: Check the plotted values.
```
>>> assert get_hash('X' in ax_2.get_xlabel().upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check x-axis label.
```
>>> assert get_hash('PROBABILITY' in ax_2.get_ylabel().upper()) == 'f827cf462f62848df37c5e1e94a4da74'
```
failure_message: Check y-axis label.
```
>>> keywords = ['PMF', 'OF', 'X']
>>> assert all((word in ax_2.get_title().upper() for word in keywords))
```
failure_message: Check plot title.


Tests for q5.0 (points: 0.5)
```
>>> assert isinstance(q5_0, tuple)
>>> assert all((isinstance(option, str) for option in q5_0))
>>> assert all((len(option) == 1 for option in q5_0))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.
```
>>> assert any(('800618943025315f869e4e1f09471012' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q5_0))
```
failure_message: Incorrect answer.


Tests for q5.1 (points: 0.5)
```
>>> assert isinstance(q5_1, tuple)
>>> assert all((isinstance(option, str) for option in q5_1))
>>> assert all((len(option) == 1 for option in q5_1))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.


Tests for q5.2 (points: 0.5)
```
>>> assert get_hash(len(CDF_X)) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(sum(CDF_X)) == 'aa756336a7580553c54be76c83fcdc25'
```
failure_message: Incorrect answer.


Tests for q6.0 (points: 0.5)
```
>>> assert type(fig_3).__name__ == 'Figure'
>>> assert type(ax_3).__name__ == 'Axes'
```
failure_message: Make sure to create a figure with 1 axes and assign them to "fig_3" and "ax_3".
```
>>> assert isinstance(ax_3.lines[0], Line2D)
```
failure_message: Make sure you are plotting to ax_3 using step().
```
>>> assert get_hash(sum(ax_3.lines[0].get_xdata())) == '33e75ff09dd601bbe69f351039152189'
>>> assert get_hash(len(ax_3.lines[0].get_xdata())) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(sum(ax_3.lines[0].get_ydata())) == 'aa756336a7580553c54be76c83fcdc25'
>>> assert get_hash(len(ax_3.lines[0].get_ydata())) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(ax_3.lines[0].get_drawstyle()) == '476489a3f6634409f06ea1d7e6eef3fa'
```
failure_message: Check the arguments of plt.step()
```
>>> assert get_hash(q6_0) == 'f8320b26d30ab433c5a54546d21f414c'
```
failure_message: Incorrect answer. Are the jumps equal to the PMF values at each x value?


Tests for q6.1 (points: 0.5)
```
>>> assert type(fig_4).__name__ == 'Figure'
>>> assert type(ax_4).__name__ == 'Axes'
```
failure_message: Make sure to create a figure with 1 axes and assign them to "fig_4" and "ax_4".
```
>>> assert isinstance(ax_4.lines[0], Line2D)
```
failure_message: Make sure you are plotting to ax_4 using step().
```
>>> assert get_hash(sum(ax_4.lines[0].get_xdata())) == '33e75ff09dd601bbe69f351039152189'
>>> assert get_hash(len(ax_4.lines[0].get_xdata())) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(sum(ax_4.lines[0].get_ydata())) == 'aa756336a7580553c54be76c83fcdc25'
>>> assert get_hash(len(ax_4.lines[0].get_ydata())) == 'c9f0f895fb98ab9159f51fd0297e236d'
>>> assert get_hash(ax_4.lines[0].get_drawstyle()) == 'ae9b02c3ff4dccafc125288ab2f2cfd8'
```
failure_message: Check the arguments of plt.step()


Tests for q7.0 (points: 0.25)
```
>>> assert get_hash(round(q7_0, 10)) == '6780987c038ab92332702313b6c3f64a'
```
failure_message: Incorrect answer.


Tests for q7.1 (points: 0.25)
```
>>> assert get_hash(round(q7_1, 10)) == '5d6182b8169f820c3e247e91131138ea'
```
failure_message: Incorrect answer.


Tests for q7.2 (points: 0.25)
```
>>> assert get_hash(round(q7_2, 10)) == '1033ec73b013426ab1d5b8357e546714'
```
failure_message: Incorrect answer.


Tests for q7.3 (points: 0.5)
```
>>> assert get_hash(round(q7_3, 10)) == '7e4fb13bfeb75d9fa2b89864024247f2'
```
failure_message: Incorrect answer.


Tests for q8.0 (points: 0.5)
```
>>> assert get_hash(round(mean_X, 10)) == 'cd294afb08705a802b70df0a63743c16'
```
failure_message: Incorrect answer.


Tests for q8.1 (points: 0.5)
```
>>> assert get_hash(round(stdev_X, 10)) == 'de14d24886946cd9a7bd80becf861ec0'
```
failure_message: Incorrect answer.


Tests for q8.2 (points: 0.5)
```
>>> assert get_hash(round(skew_X, 10)) == '3dbfc293c1f8b5c0206e86eefef2e028'
```
failure_message: Incorrect answer.


Tests for q8.3 (points: 0.25)
```
>>> assert len(q8_3) == 1
>>> assert isinstance(q8_3, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q8_3.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q8.4 (points: 0.5)
```
>>> assert get_hash(round(mean_month, 10)) == '3af4a86e137732480173bbc9f21f2e7f'
```
failure_message: Incorrect answer.


Tests for q8.5 (points: 0.5)
```
>>> assert get_hash(round(stdev_month, 10)) == '9f55f7a65ca90890c9f0aea02f6657cb'
```
failure_message: Incorrect answer.
