```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# CE 93: Lab Assignment 01

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to revisit fundamental Python commands. As outlined in the syllabus, Data 8, Engin 7, or equivalent are prerequisites for CE 93. More specifically, basic programming experience is required to succeed in this course. While CE 93 will extend your Python proficiency, we won't cover the basics, assuming that students already possess this foundational knowledge. If you find it challenging to provide correct answers to more than 70% of the lab questions, it is advisable to wait until you have completed Data 8 (or an equivalent course) before enrolling in CE 93.

## Instructions 

**Watch [this tutorial](https://drive.google.com/file/d/19fYhFW8V6utgTT-grDqwIYQycPMdc8MR/view?usp=sharing) before you start the assignment. This step by step tutorial will show you how to work on the lab assignment and correctly submit your work to Gradescope.**

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Any part listed as a "<font color='red'>**Question**</font>" should be answered to receive credit.

**Please save your work after every question!**

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

import math
import numpy as np
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import resources
```

## Autograding and Gradescope

We will be using an automatic grader for lab assignments. Each question in a lab assignment will have an **Answer Cell** followed by a **Testing Cell**. You will put gradable code into the **Answer Cell** as instructed in the lab, and then test your code by executing the **Testing Cell**. The output of the testing cell will let you know whether you have a mistake in the graded cell.

### Example Question

Given two variables, `m=1` and `n=2`, compute their sum and assign it to variable `o`. 

+ The correct solution to this exercise is to replace the `...` below with `m + n`.  Try that out, and confirm that the testing cell marks your answer as correct. 
+ Also try an *incorrect* answer, by replacing the `...` with `-2`. Notice what happens when you run the testing cell.

It is important that you save your work regularly, preferably after answering each question. You can save your work by clicking on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. Failing to save your work may result in incomplete submissions and loss of points.

#### Answer Cell:

```python
m = 1
n = 2
o = ...
```

#### Testing Cell:

```python
grader.check("q0")
```

### Python Basics

<font color='red'>**Question 1.0.**</font> What does a hashtag (`#`) represent in Python syntax? Assign your answer to the variable `q1_0` as a string. (0.25 pts)

A. Assignment \
B. Function \
C. Comment \
D. Selection \
E. Iteration

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
For example, if you selected `A. Assignment`, then simply replace `...` in the cell below by `"A"`.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q1_0 = ...
q1_0
```

```python
grader.check("q1.0")
```

### Mathematical Operations

An arithmetic **operation** is either addition, subtraction, multiplication, division, or exponentiation (i.e., power). An arithmetic **operator** is a symbol that Python has reserved to perform one of the aforementioned operations.

| Operation       | Mathematical Notation | Python Operator | Python Example | Output    |
|:----------------|:---------------------:|:---------------:|:---------------|:----------|
| Addition        | $a+b$                 | `+`             | `3 + 1`        | `5`       |
| Subtraction     | $a-b$                 | `-`             | `3 - 2`        | `1`       |
| Multiplication  | $a\times b$           | `*`             | `3 * 2`        | `6`       |
| Division        | $\dfrac{a}{b}$        | `/`             | `3 / 2`        | `1.5`     |
| Exponentiation  | $a^b$                 | `**`            | `3 ** 2`       | `9`       |

Python obeys the same order of operations that you learned in school: powers are executed before multiplication and division, which are executed before addition and subtraction. Parentheses, `()`, can also be used in Python to group together smaller expressions and supersede the standard order of operations. The order in which Python executes these operations is shown in the table below.

| Precedence      | Operation                      | Rule                                       | 
|:----------------|:------------------------------ | :----------------------------------------- |
| First (Highest) | Parentheses: `( )`             | Evaluated starting with the innermost pair |
| Second          | Exponentiation: `**`           | Evaluated from left to right               | 
| Third           | Unary positive, negative: `+x`, `-x` | Evaluated from left to right         |
| Fourth          | Multiplication, division: `*`, `/`   | Evaluated from left to right         |
| Fifth           | Addition, subtraction: `+`, `-`      | Evaluated from left to right         |

<font color='red'>**Question 2.0.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q2_0` as a string. (0.25 pts)

```python
2 * 3 ** 2
```

A. 9 \
B. 12 \
C. 18 \
D. 36 \
E. `Error`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q2_0 = ...
q2_0
```

```python
grader.check("q2.0")
```

<font color='red'>**Question 2.1.**</font> Construct a one-line Python expression that computes the value of the following mathematical expression. Assign the expression to the variable `q2_1`. Your answer should be a mathematical expression, not the value of the expression. (0.25 pts)

$$3+\frac{1}{3}$$

```python
# ANSWER CELL
q2_1 = ...
q2_1
```

```python
grader.check("q2.1")
```

Python has many basic trigonometric, logarithmic, and other arithmetic functions like `sin()`, `cos()`, `tan()`, `asin()`, `acos()`, `atan()`, `exp()`, `log()`, `log10()` and `sqrt()` stored in a module called **`math`**. You can read about all the functions provided by the **`math`** module [here](https://docs.python.org/3/library/math.html).

The way we use these mathematical functions (and in general, any function from a module) is `module.function(x)`, where:
* `module` is the module's name or its given shortened name; e.g., `math`
* `.function` is the function available within that module that you want to use; e.g., `.log10`
* `(x)` includes inside the parentheses the input(s) that your are trying to evaluate using the function; e.g., `(100)`.

**Example:** 

```python
>>> math.log10(100)
2.0
```

<font color='red'>**Question 2.2.**</font>  Define variables $x=2$ and $y=5$. Then, construct a one-line Python expression that computes the value of the following mathematical expression. Assign the expression to the variable `q2_2`. (0.5 pts)

$$e^{x^y}$$

Your Python expression should give appropriate values regardless of the values of $x$ and $y$ (i.e. you should use the variable names instead of their numerical values).

*Hint:* Take a look at `math`'s exponential function documentation [here](https://docs.python.org/3/library/math.html#power-and-logarithmic-functions).

```python
# ANSWER CELL
x = ...
y = ...
q2_2 = ...
q2_2
```

```python
grader.check("q2.2")
```

### Variables and Assignment

**Variables** are used to store values in Python using an *assignment* statement. In an *assignment* statement, a name is followed by `=`, which is followed by any expression. That's why <code>=</code> is referred to as the assignment operator.

**Example:**

```python
a = 5
b = a
```

The value of the expression to the right of `=` is *assigned* to the variable name to the left of `=`. So, in the example above, the value 5 is assigned to the variable `a`. After executing this line, the variable `a` behaves like the value 5 in subsequent calculations, until the variable is changed or deleted.

When a variable is assigned, it has no memory of how it was assigned. We assigned variable <code>b = a</code>. However, changing the value of <code>a</code> will not automatically change the value of <code>b</code>. To change the value of <code>b</code>, you have to **reassign it**.

Note that one way to display values in Python is using the `print()` function, which takes as input some expression and prints its value.

<font color='red'>**Question 3.0.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_0` as a string. (0.25 pts)

```python
a = [1, 2, 3]
b = a
a = [4, 5, 6]
print(b)
```

A. `[1, 2, 3]` \
B. `b` \
C. `[4, 5, 6]` \
D. `a` \
E. No output

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_0 = ...
q3_0
```

```python
grader.check("q3.0")
```

<font color='red'>**Question 3.1.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_1` as a string. (0.25 pts)

```python
a = 10
b = a
a = 15
```

A. `10` \
B. `15` \
C. `b` \
D. `a` \
E. No output

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_1 = ...
q3_1
```

```python
grader.check("q3.1")
```

The statement <code>a = 5</code> is not asserting that <code>a</code> is already equal to 5, as we might expect in mathematical notation. Rather, this line of code defines the variable <code>a</code> to hereafter have the value 5.

Although it is perfectly valid to say $5 = a$ in mathematics, `5 = a` will raise an error in Python.

When trying to run `5 = a`, you are telling Python to define a variable named `5` and assign it the value of the variable `a`. However, `5` is not a valid variable name. That's why Python raises a `SyntaxError`. There are some restrictions on the names variables can take in Python, which are listed hereafter:
* Variable names must start with a letter or underscore `_` but cannot start with a number
* Variables can only contain alphanumeric characters (letters and numbers) as well as underscores `_` 
* Numbers are allowed after the first character of the variable name
* No spaces or math symbols/punctuation within a variable name are permitted
* Variable names are case-sensitive (e.g., `a` and `A` will be considered different variables)

So, variable names must start with a letter, but can contain both letters and numbers after the first character. A name cannot contain a space; instead, it is common to use an underscore `_` character to replace each space. For example, `1dice_roll = 5` will raise a `SyntaxError` because a name cannot start with a number. Instead, you can write `dice_roll_1 = 5`, and then `dice_roll_1` will be assigned the value $5$.

<font color='red'>**Question 3.2.**</font>  What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q3_2` as a string. (0.25 pts)

```python
2 = 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q3_2 = ...
q3_2
```

```python
grader.check("q3.2")
```

<font color='red'>**Question 3.3.**</font> Which of the following is (are) NOT valid variable name(s)? Assign ALL that apply to the variable `q3_3`. (0.5 pts)

A. `Data` \
B. `Data1` \
C. `1Data` \
D. `^_^` \
E. `@potus`\
F. `Data 1`\
G. `Data_1`

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q3_3 = ...
q3_3
```

```python
grader.check("q3.3")
```

### Logical Expressions and Operators

**Logical expressions** are statements that can either be true or false. Logical expressions are used to pose questions to Python. For example, running `2 < 3` is equivalent to asking , "*Is 2 less than 3?*" Since this statement is true, Python will return the value `True`. Otherwise, if the statement is false, Python will return `False`. The returned value is a data type which is known as **boolean** (or `bool`), which has the built-in values `True` and `False`.

Logical expressions can be applied on values, variables, or both. Python includes a variety of comparison operators (operators to perform comparisons). 

| Comparison         | Python Operator | True example | False Example |
|:-------------------|:---------------:|:-------------|:--------------|
| Less than          | `<`             | `2 < 3`      | `2 < 2`       |
| Greater than       | `>`             | `3 > 2`      | `3 > 3`       |
| Less than or equal | `<=`            | `2 <= 2`     | `3 <= 2`      |
| Greater or equal   | `>=`            | `3 >= 3`     | `2 >= 3`      |
| Equal              | `==`            | `3 == 3`     | `3 == 2`      |
| Not equal          | `!=`            | `3 != 2`     | `2 != 2`      |

<font color='red'>**Question 4.0.**</font> Recall that you have previously defined `x = 2`. What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_0` as a string. (0.25 pts)

```python
x != 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_0 = ...
q4_0
```

```python
grader.check("q4.0")
```

Two consecutive equal marks (`==`) are used to check whether two expressions have the same value. For example, if we define `dice_roll_1 = 5` and then run `dice_roll_1 == 5`, this will return `True` because we have assigned the value $5$ to `dice_roll_1`.

<font color='red'>**Question 4.1.**</font>  What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_1` as a string. (0.25 pts)

```python
x == 3
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_1 = ...
q4_1
```

```python
grader.check("q4.1")
```

In Python and many other programming languages, `True` is equivalent to 1 and `False` is equivalent to 0. So, you can perform arithmetic operations on `bool` data types. Because `True` is equal to 1 and `False` is equal to 0, adding logical expressions together is a quick way to count the number of `True` expressions out of many logical expressions.

<font color='red'>**Question 4.2.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. (0.25 pts)

```python
(5 > 4) + 2 * (3 != 10) - (7 < 10)
```

Assign your answer to the variable `q4_2` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q4_2 = ...
q4_2
```

```python
grader.check("q4.2")
```

In most of the previous examples, we did one comparison at a time. **Logical operators** allow us to combine together multiple comparisons. There are three common logical operators: AND, OR, NOT. Operator symbols vary with different programming languages. In Python they are `and`, `or`, and `not`.

| Operator | Description | Examples |
| :---:    | :---:       | :---:    | 
| `and`    | `True` if **all** expressions are `True` <br> `False` otherwise| `(5 > 4) and (5 > 3)`<br>`(5 > 4) and (5 < 3)`|
| `or`     | `True` if **at least one** expression is `True` <br> `False` otherwise| `(5 > 4) or (5 < 3)`<br>`(5 < 4) or (5 < 3)`|
| `not`    | Reverses the result, `True` if the expression is `False` <br> `False` if the expression is `True` |`not(5 < 4)`<br>`not(5 > 4)`| 

<font color='red'>**Question 4.3.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q4_3` as a string. (0.25 pts)

```python
(5 > 4 or 2 == 3)  and  (3 != 10 and 7 < 10)
```

A. `True` \
B. `False` \
C. `NaN` \
D. `NA` \
E. `SyntaxError`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q4_3 = ...
q4_3
```

```python
grader.check("q4.3")
```

### Strings

A string is a sequence of characters, such as `"Hello World"`. Strings are surrounded by either single `' '` or double quotation marks `" "`. A string can include letters, digits, punctuation, spaces, and other valid symbols. They can represent anything from a single character to a word, a sentence, or even an entire book! Strings can also be assigned to variables.

Numbers can also be expressed as strings. For example, `z = "93"` means that `z` is assigned the string 93 not the number 93.

Strings can be concatenated (glued together) with the `+` operator. This allows you to combine multiple strings into a single string. So, if we define `z = '93'` and then run `"Welcome " + "to " + "CIVENG" + z`, this will return the string `"Welcome to CIVENG93"`. Notice how spaces were included after "Welcome" and after "to", but not after "CIVENG".

<font color='red'>**Question 5.0.**</font> What is the result of this Python expression? Feel free to run it to check your answer, but try to think it through first. Assign your answer to the variable `q5_0` as a string. (0.25 pts)

```python
'2' + '222'
```

A. `2222` \
B. `224` \
C. `222` \
D. `'2222'` \
E. `'224'`

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q5_0 = ...
q5_0
```

```python
grader.check("q5.0")
```

In Python, if you try to concatenate a string with a numeric value, you will get a `TypeError`.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> course_code + course_number
```
```
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)

TypeError: can only concatenate str (not "int") to str
```

The correct way to do it is to convert the numeric value to string first, and then concatenate them. This can be achieved using the built-in function `str()`. For example `str(1)` will convert the number 1, to the string '1'. This is particularly useful when you have calculated a numeric variable and want to print it out with a string. In this case, you can convert the numeric variable to a string and then concatenate it with the string you are trying to print out.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> course_code + str(course_number)
'CE93'
```

A more elegant way to combine numeric data with strings is known as formatted strings. Prefix the string with the letter `f` and then any variables, which could be strings, numeric, or other are included between `{ }`.

```python
>>> course_code = 'CE'
>>> course_number = 93
>>> f'Welcome to {course_code}{course_number}!'
'Welcome to CE93!'
```

<font color='red'>**Question 5.1.**</font> Consider the Python code below. Which of the options below could replace `...` to get the following output: `'The average rainfall is 0.925 inch'`? Assign ALL that apply to the variable `q5_1`. (0.5 pts)

```python
rain_day1 = 1.1
rain_day2 = 0.75
...
```

A. `'The average rainfall is ' + str((rain_day1 + rain_day2) / 2) + ' inch'` \
B. `'The average rainfall is ' + (rain_day1 + rain_day2) / 2 + ' inch'` \
C. `'The average rainfall is {(rain_day1 + rain_day2) / 2} inch'` \
D. `f'The average rainfall is {(rain_day1 + rain_day2) / 2} inch'` \
E. `f'The average rainfall is (rain_day1 + rain_day2) / 2 inch'`

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.
Assign your answer to the given variable. Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q5_1 = ...
q5_1
```

```python
grader.check("q5.1")
```

### Lists and NumPy Arrays

A `list` is a more versatile data structure that is a collection of one or more items. Lists can hold various data types or a combination of different data types, such as strings and numeric data. Items in a list are enclosed within square brackets `[]` and separated by commas `,`.

**Examples:**

```python
>>> numeric_list = [1, 2, 3]
>>> numeric_list
[1, 2, 3]

>>> text_list = ['Welcome', 'to', 'CE93!']
>>> text_list
['Welcome', 'to', 'CE93!']

>>> mixed_list = ['Welcome', 'to', 'CE', 93, '!']
>>> mixed_list
['Welcome', 'to', 'CE', 93, '!']
```
Lists are fundamental in Python and serve as a powerful tool for storing and manipulating collections of data.

We can use indexing to access an element from a list. To access an element from a list, we use a pair of square brackets. Indexing in Python starts at 0, which means that the first element has an index of 0, the second element has an index of 1, and so on.

```python
>>> numeric_list = [12, 52, 32]
>>> numeric_list[0]
12
```

You can also use negative indexes. Negative indexes simply means counting from the end of the list. So, -1 is the index of the last element, -2 is the index of the second-to-last element, and so on.

```python
>>> numeric_list = [12, 52, 32]
>>> numeric_list[-1]
32
```

<font color='red'>**Question 6.0.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. (0.5 pts)

```python
a = [383, 165, 273, 112, 937, 11, 64, 112, 455]
print(a[1])
```

Assign your answer to the variable `q6_0` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q6_0 = ...
q6_0
```

```python
grader.check("q6.0")
```

There are essential built-in functions that apply to data structures in Python. The operations in the table below are widely supported by most data structure types and form the foundation of data manipulation in Python.

In the table, `s` is a data structure, such as a list, and `i` is an integer.

| Operation      | Result                                                                      | 
| :------------- | :-------------------------------------------------------------------------- |
| `s[i]`         | Index and retrieve the item at position `i` of `s`                          |
| `len(s)`       | Return the length (number of items) in `s`                                  |
| `min(s)`       | Return the smallest or lowest item in `s`                                   |
| `max(s)`       | Return the largest or highest item in `s`                                   |

<font color='red'>**Question 6.1.**</font> Write a Python expression that extracts the maximum element in `a` and the minimum element in `a` and then divides their sum by the length of `a`. Assign this expression to `q6_1`. (0.5 pts)

```python
# ANSWER CELL
a = [383, 165, 273, 112, 937, 11, 64, 112, 455]
q6_1 = ...
q6_1
```

```python
grader.check("q6.1")
```

Another essential data structure for scientific and numerical applications is the NumPy array, which is part of the NumPy module. NumPy is short for "Numerical Python". 

In order to use NumPy, we need to import it first. A conventional way to import it is to use `np` as a shortened name: `import numpy as np`.

To define an array in Python, you could use the `np.array()` function. The items are put inside square brackets `np.array([...])` or a second set of parentheses `np.array((...))`, and are separated by commas `,`.

**Example:**

```python
>>> w = np.array([1, 4, 3, 12])
```

Arrays can have multiple rows and columns. In NumPy, attributes are properties of NumPy arrays that provide information about the array's shape, size, dimension, and so on. There are many attributes, which you can read about in the documentation [here](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.html). Some of the common attributes are:
* [`array_name.shape`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.shape.html): returns the shape of `array_name`
> returns `(n_rows, n_columns)`, where the first element is the number of rows in `array_name` and the second element is the number of columns in `array_name` 
* [`array_name.size`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.size.html#numpy.ndarray.size): returns the total number of elements in `array_name` 
> number of rows times number of columns
* [`array_name.ndim`](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.ndim.html#numpy.ndarray.ndim): returns the number of dimensions of `array_name`

In addition to `np.array()`, there are a number of other functions for creating arrays. Very often we would like to define arrays that have a pattern. For instance, we may wish to create the array `days = [1, 2, 3, ..., 365]`. It would be very cumbersome to type all the values of `days` into Python. For generating arrays that have a pattern, there are several helpful functions in NumPy.

For generating arrays that are in **order and evenly spaced**, it is useful to use the `np.arange()` function in NumPy:

```python
np.arange(start, stop, step)
```

where:

* `start`: a number specifying at which value to start the array. This is an optional argument, and thus, if not specified, it will take the default value 0.
* `stop`: a number specifying at which value to stop the array (excluded). This is a required argument.
* `step`: a number specifying the increment in the values of the array (`step` $\neq 0$). This is an optional argument, and thus, if not specified, it will take the default value 1. The `step` can be positive or negative and the value in index `i` will simply be `start + step*i`

**Example:**

```python
>>> days = np.arange(1, 366)
```

For generating arrays that have **a certain number of evenly spaced points between a start and end value**, it is useful to use the `np.linspace()` function in NumPy:

```python
np.linspace(start, stop, num)
```

where:

* `start`: a number specifying at which value to start the array. Here, this is a required argument.
* `stop`: a number specifying at which value to stop the array. This is by default **included** in the array, unlike `np.arange()`. This is a required argument.
* `num`: an integer number specifying how many values to generate between `start` and `stop`, including `start` and `stop`. This is an optional argument, and thus, if not specified, it will take the default value 50.

**Example:**

```python
>>> days = np.linspace(1, 365, 365)
```

Similar to strings and lists, you can you can index a NumPy array.

<font color='red'>**Question 6.2.**</font> What will be displayed after running this Python code? Feel free to run it to check your answer, but try to think it through first. (1 pt)

```python
x1 = np.arange(10, 20, 1)
x2 = np.linspace(0, 10, 3)
print(x1[-1] + x2[1])
```

Assign your answer to the variable `q6_2` in the next cell. Do NOT put quotes around your answer. Your answer should be a numeric value.

```python
# ANSWER CELL
q6_2 = ...
q6_2
```

```python
grader.check("q6.2")
```

The questions below are related to course policies on assignment submission, grading, Ed, etc. All information can be found in the syllabus on bCourses. 

Reading the syllabus will save you time!

These questions are graded and you must answer them correctly to receive credit.

<font color='red'>**Question 7.0.**</font> You have a general question regarding a homework or lab assignment. Who do you contact? Assign ALL that apply to the variable `q7_0`. (0.25 pts)

A. Email the instructor \
B. Post a question on Ed  \
C. Copy the solution  \
D. ce93@berkeley.edu \
E. Go to office hours

Answer in the next cell. Add each selected choice as a string and separate each two answer choices by a comma. For example, if you want to select `"A"` and `"B"`, your answer should be `"A", "B"`.\
Assign your answer to the given variable.
Remember to put quotes around each answer choice.

```python
# ANSWER CELL
q7_0 = ...
q7_0
```

```python
grader.check("q7.0")
```

<font color='red'>**Question 7.1.**</font> Homework and lab assignments are due by 11:59 PM. Which of the following is true about late homework and lab submissions? Assign your answer to the variable `q7_1` as a string. (0.5 pts)

A. No late submissions are allowed \
B. Assignments will be accepted for 24 hours after the deadline but you have to submit an extension request \
C. Assignments will be accepted for 24 hours after the deadline without the need to request an extension \
D. Late assignments can be requested by emailing the course staff

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. \
Assign your answer to the given variable.
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q7_1 = ...
q7_1
```

```python
grader.check("q7.1")
```

<font color='red'>**Question 7.2.**</font> If you need more time beyond the default 24 hour extension, what should you do? (0.25 pts)

A. No further late submissions are allowed \
B. Email the teaching staff asking for an extension \
C. Submit the extension form on bCourses before the original due date \
D. Submit the extension form on bCourses anytime

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. Assign your answer to the variable `q7_2` as a string.

```python
# ANSWER CELL
q7_2 = ...
q7_2
```

```python
grader.check("q7.2")
```

<font color='red'>**Question 7.3.**</font> To ensure everyone is familiar with the lab/homework extension process, all students are required to complete the Lab 01 Extension Request form, even if you do not need an extension.

Instructions:

* Go to bCourses and click on Support & Forms from the home page
* Click on the Lab 01 Extension under Request a Lab Extension
* Open the form and follow the prompts
* If you do not need an extension, simply select the No option on the second question
* Submit the form

This is just a way for everyone to get comfortable with the system.

Please use your Berkeley email to access the form. At the end of the form, there will be a secret word, which you must enter in the box below. Be sure to enclose the secret word in quotation marks (e.g., 'hello'), as this indicates it is a String type. (0.25 pts)

Submitting an answer without completing the form will be considered academic misconduct. Therefore, please fill out the form and do not attempt to acquire the secret word without doing so.

```python
# ANSWER CELL
q7_3 = ...
q7_3
```

```python
grader.check("q7.3")
```

<font color='red'>**Question 7.4.**</font> Which of the following is(are) true regarding academic misconduct? (0.5 pts)

A. The penalty for assignment misconduct is $-$10 on that assignment for all involved, and which cannot be dropped \
B. A repeat offense on assignments will result in an F in the course \
C. Anyone involved in misconduct on any exam will receive an F in the course \
D. Any misconduct incident could be reported to the Center for Student Conduct \
E. If I need help, I can go to office hours, post on Ed, and benefit from the course support. I do not need to cheat! \
F. All of the above

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. Assign your answer to the variable `q7_4` as a string.

```python
# ANSWER CELL
q7_4 = ...
q7_4
```

```python
grader.check("q7.4")
```

<font color='red'>**Question 7.5.**</font> Based on what criteria will homework assignments be graded? (0.5 pts)

A. Correctness \
B. Completeness and correctness \
C. Completeness, self-assessment, and professionalism \
D. Accuracy

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. Assign your answer to the variable `q7_5` as a string.

```python
# ANSWER CELL
q7_5 = ...
q7_5
```

```python
grader.check("q7.5")
```

<font color='red'>**Question 7.6.**</font> Which of the following is(are) true regarding iClicker questions? (0.5 pts)

A. iClicker questions are incorporated into each lecture and count toward the grade \
B. These questions are graded for completeness, not correctness \
C. If I miss a lecture or have issues accessing iClicker, I can complete a short online quiz on bCourses, due the same day as the lecture \
D. If I respond to both the iClicker questions and the bCourses quiz, my highest score for that lecture will be kept \
E. To earn a perfect iClicker score, I need a total of 30 points. Any points beyond 30 will count as bonus points, up to +9. \
F. There will be no makeups for iClicker questions, but bonus points can be earned for answering questions correctly \
G. All of the above

Answer in the next cell. Your answer should be a string, e.g., `"A"`, `"B"`, etc. Assign your answer to the variable `q7_6` as a string.

```python
# ANSWER CELL
q7_6 = ...
q7_6
```

```python
grader.check("q7.6")
```

<font color='red'>**Question 7.7.**</font> To ensure office hours work for as many students as possible, please fill out [this form](https://forms.gle/kdDL8UZ3hbdaZeE88) with your availability. Your responses will help us determine the best times to hold office hours this semester.  

Please use your Berkeley email to access the form. At the end of the form, there will be a secret word, which you must enter in the box below. Be sure to enclose the secret word in quotation marks (e.g., 'hello'), as this indicates it is a String type. (1 pt)

Submitting an answer without completing the form will be considered academic misconduct. Therefore, please fill out the form and do not attempt to acquire the secret word without doing so.

```python
# ANSWER CELL
secret_word = ...
```

```python
grader.check("q7.7")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Finally, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing your first lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q0 (points: 0)
```
>>> assert get_hash(o) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q1.0 (points: 0.25)
```
>>> assert len(q1_0) == 1
>>> assert isinstance(q1_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q1_0.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: Incorrect answer.


Tests for q2.0 (points: 0.25)
```
>>> assert len(q2_0) == 1
>>> assert isinstance(q2_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q2_0.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
failure_message: Incorrect answer.


Tests for q2.1 (points: 0.25)
```
>>> assert get_hash(round(q2_1, 10)) == '012bd0771172d790b0c88107c9743442'
```
failure_message: Incorrect answer.


Tests for q2.2 (points: 0.5)
```
>>> assert get_hash(round(x, 1)) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(round(y, 1)) == 'e4da3b7fbbce2345d7772b0674a318d5'
```
failure_message: Incorrect x and/or y. Make sure x and y are defined based on the given values.
```
>>> assert get_hash(round(q2_2, 2)) == '027e97349ab774d836e6bde002c96d80'
```
failure_message: Incorrect answer for q2_2.


Tests for q3.0 (points: 0.25)
```
>>> assert len(q3_0) == 1
>>> assert isinstance(q3_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q3_0.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q3.1 (points: 0.25)
```
>>> assert len(q3_1) == 1
>>> assert isinstance(q3_1, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q3_1.upper()) == '3a3ea00cfc35332cedf6e5e9a32e94da'
```
failure_message: Incorrect answer.


Tests for q3.2 (points: 0.25)
```
>>> assert len(q3_2) == 1
>>> assert isinstance(q3_2, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q3_2.upper()) == '3a3ea00cfc35332cedf6e5e9a32e94da'
```
failure_message: Incorrect answer.


Tests for q3.3 (points: None)
```
>>> assert isinstance(q3_3, tuple)
>>> assert all((isinstance(option, str) for option in q3_3))
>>> assert all((len(option) == 1 for option in q3_3))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert not any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert not any(('dfcf28d0734569a6a693bc8194de62bf' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.
```
>>> assert any(('800618943025315f869e4e1f09471012' == get_hash(option.upper()) for option in q3_3))
```
failure_message: Incorrect answer.


Tests for q4.0 (points: 0.25)
```
>>> assert len(q4_0) == 1
>>> assert isinstance(q4_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q4_0.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q4.1 (points: 0.25)
```
>>> assert len(q4_1) == 1
>>> assert isinstance(q4_1, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q4_1.upper()) == '9d5ed678fe57bcca610140957afab571'
```
failure_message: Incorrect answer.


Tests for q4.2 (points: 0.25)
```
>>> assert isinstance(q4_2, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
success_message: Correct answer format.
```
>>> assert get_hash(int(q4_2)) == 'c81e728d9d4c2f636f067f89cc14862c'
```
failure_message: Incorrect answer.


Tests for q4.3 (points: 0.25)
```
>>> assert len(q4_3) == 1
>>> assert isinstance(q4_3, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q4_3.upper()) == '7fc56270e7a70fa81a5935b72eacbe29'
```
failure_message: Incorrect answer.


Tests for q5.0 (points: 0.25)
```
>>> assert len(q5_0) == 1
>>> assert isinstance(q5_0, str)
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes.
success_message: Correct answer format.
```
>>> assert get_hash(q5_0.upper()) == 'f623e75af30e62bbd73d6df5b50bb7b5'
```
failure_message: Incorrect answer.


Tests for q5.1 (points: None)
```
>>> assert isinstance(q5_1, tuple)
>>> assert all((isinstance(option, str) for option in q5_1))
>>> assert all((len(option) == 1 for option in q5_1))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert not any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert not any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.
```
>>> assert not any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q5_1))
```
failure_message: Incorrect answer.


Tests for q6.0 (points: 0.5)
```
>>> assert isinstance(q6_0, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
success_message: Correct answer format.
```
>>> assert get_hash(int(q6_0)) == '9766527f2b5d3e95d4a733fcfb77bd7e'
```
failure_message: Incorrect answer. Indexing in Python starts at 0.


Tests for q6.1 (points: 0.5)
```
>>> assert isinstance(q6_1, float)
```
failure_message: Incorrect answer format. Your answer should be a float.
success_message: Correct answer format.
```
>>> assert get_hash(round(q6_1, 10)) == '6d025b8f42889608be47f9567281fdd8'
```
failure_message: Incorrect answer. The entire sum should be divided by the length.


Tests for q6.2 (points: 1)
```
>>> assert isinstance(q6_2, int)
```
failure_message: Incorrect answer format. Your answer should be an integer.
success_message: Correct answer format.
```
>>> assert get_hash(int(q6_2)) == '1ff1de774005f8da13f42943881c655f'
```
failure_message: Incorrect answer. Indexing in Python starts at 0.


Tests for q7.0 (points: None)
```
>>> assert isinstance(q7_0, tuple)
>>> assert all((isinstance(option, str) for option in q7_0))
>>> assert all((len(option) == 1 for option in q7_0))
```
failure_message: Incorrect answer format. Make sure you only have the letter of your answer choice in quotes and answers are separated by commas.
success_message: Correct answer format.
```
>>> assert any(('9d5ed678fe57bcca610140957afab571' == get_hash(option.upper()) for option in q7_0))
```
failure_message: Correct! For most questions, Ed is the right place to ask them. If you send a general question through email, we will direct you to post it on Ed.
```
>>> assert any(('3a3ea00cfc35332cedf6e5e9a32e94da' == get_hash(option.upper()) for option in q7_0))
```
failure_message: Correct! There are several office hours throughout the week. Check bCourses for the schedule.
```
>>> assert not any(('7fc56270e7a70fa81a5935b72eacbe29' == get_hash(option.upper()) for option in q7_0))
```
failure_message: Incorrect answer. Emailing the instructor will generally lead to slower response.
```
>>> assert not any(('0d61f8370cad1d412f80b84d143e1257' == get_hash(option.upper()) for option in q7_0))
```
failure_message: Incorrect answer.
```
>>> assert not any(('f623e75af30e62bbd73d6df5b50bb7b5' == get_hash(option.upper()) for option in q7_0))
```
failure_message: Incorrect answer. ce93@berkeley.edu should only be used for personal inquiries and not for general questions.


Tests for q7.1 (points: 0.5)
```
>>> assert get_hash(q7_1.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
success_message: Correct! Please do not email us to request short extensions, as 24 hours are granted by default.


Tests for q7.2 (points: 0.25)
```
>>> assert get_hash(q7_2.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
success_message: Correct! Long assignment extension requests (> 24 hours) should be submitted before the original due date, and NOT within the 24-hour grace period.


Tests for q7.3 (points: 0.25)
```
>>> assert get_hash(q7_3.upper()) == 'e7c374316a119ab9c10eeb6c11e6fa97'
```


Tests for q7.4 (points: 0.5)
```
>>> assert get_hash(q7_4.upper()) == '800618943025315f869e4e1f09471012'
```
success_message: Correct! Deadlines can be stressful. If you feel that you need to engage in academic misconduct to meet a deadline, please reach out to us. We care that you learn the material and we are here to support you.


Tests for q7.5 (points: 0.5)
```
>>> assert get_hash(q7_5.upper()) == '0d61f8370cad1d412f80b84d143e1257'
```
success_message: Correct! A guide will be provided for each homework assignment to assist you whenever you get stuck. Thus, assignments are not graded for correctness.


Tests for q7.6 (points: 0.5)
```
>>> assert get_hash(q7_6.upper()) == 'dfcf28d0734569a6a693bc8194de62bf'
```
success_message: Correct! Please do not email if you miss a lecture or are have issues accessing iClicker in class. You can take the online quiz instead. If you are unable to take the online quiz, you can earn bonus points to make up for the missed lecture by answering questions correctly.


Tests for q7.7 (points: 1)
```
>>> assert get_hash(secret_word.upper()) == 'f7c65108a10ec091e16add7acbe4e771'
```
